# Math 232 Computing Assignment 3

```python
import numpy as np
import scipy.linalg as la   # package with addtional linear algebra tool: for example matrix inverse
from rref import *
```

## Some python essentials:

### Matrix Inverse

To compute the inverse of a matrix we'll need to load in the `scipy.linalg` library. This was already done above with the line:

    import scipy.linalg as la

Then we use the command `la.inv` to compute an inverse of a matrix.    

```python
# Example
B = np.array([[1,1,2],
              [1,2,2],
              [2,4,3]])
Binv = la.inv(B)
print(Binv)
```

Let's check it is the inverse of $B$. This means we want to verify $B B^{-1} = I$.

```python
B@Binv
```

### Matrix Entry Extraction
(It isn't absolutely necessary to know this, but it is helpful to know how to pick out a sub-matrix from a larger matrix.)

Suppose we have a matrix stored in memory, for example 
$ A = \begin{bmatrix}1 & 2 & 3\\ 4 & 5 & 6\end{bmatrix}. 
%\bvec{v} = \cvecthree{1}{2}{3}, 
$
To select the entry in row 2, column 2 this would be the entry with row index 1 and column index 1 (Python starts counting at $0$). To select that entry we use `A[1,1]`. 

```python
A = np.array([[1,2,3],
              [4,5,6]])
A[1,1]
```

To select the bottom right entry (the 6) this is in row with index 1 and column with index 2, so it is `A[1,2]`.<br />
**Hint:** Just begin counting at 0 when counting down rows or across columns.

```python
A[1,2]
```

We can use the colon notation `:` to select a sub-matrix consisting of rows with index $i$ to $j-1$, and columns with index $r$ to $s-1$: to do this use `A[i:j,r:s]`.

```python
A = np.array([[1, 2, 3, 4],
              [5, 6, 7, 8],
              [9, 10, 11, 12]])
A[0:2,1:4]  # extract rows with index 0,1, and columns with index 1,2,3
```

To select **all** rows or columns use `:` in the appropriate spot.
For example, to extract column with index 2 from the matrix (that is all row entries in that column) use `A[:,2]`.

```python
A[:,2]  # selects all row entries in the column with index 2
```

This technique of extracting sub-matrices of a larger matrix is called *array slicing*. You can find more about it here: https://www.w3schools.com/python/numpy/numpy_array_slicing.asp

### For Loop

A *for loop* is a structure that allows you to execute lines of code a set number of times.  The syntax of the *for loop* is:


```python
for i in range(1,5):   # loop over the numbers 1,2,3,4 (notice it starts at 1, but ends one less than 5)
    print(i**2)        # perform this line for each value of i. i**2 means i raised the power of 2.
```

See a few more examples here: https://www.w3schools.com/python/python_for_loops.asp

## <span style="color:#87CEEB">Problem 1:</span>

Consider the basis $\mathfrak{B}=\{\vec{v}_1, \vec{v}_2, \vec{v}_3, \vec{v}_4\}$ for $\mathbb{R}^4$ where 
$$
\vec{v}_1=\begin{bmatrix}0\\1\\2\\1\end{bmatrix},\qquad 
\vec{v}_2=\begin{bmatrix}1\\5\\0\\1\end{bmatrix},\qquad 
\vec{v}_3=\begin{bmatrix}2\\3\\6\\4\end{bmatrix},\qquad 
\vec{v}_4=\begin{bmatrix}1\\3\\-1\\-1\end{bmatrix}
$$

Let $S = \{ \vec{e}_1, \vec{e}_2, \vec{e}_3, \vec{e}_4\}$ be the standard basis for for $\mathbb{R}^4$.

## <span style="color:#87CEEB">Problem 1a:</span>
Calculate the matrix $P_{\mathfrak{B}\leftarrow S}$. Save the matrix with the name `PBS`. 

**Hint**: One way to do this is to form the augmented matrix $[\vec{v}_1 \vec{v}_2 \vec{v}_3 \vec{v}_4 | \vec{e}_1 \vec{e}_3 \vec{e}_3 \vec{e}_4]$ then row reduce to $[ I | P]$.

An alternative way to do this is to first determine $P_{S\leftarrow \mathfrak{B}}$ and then compute the inverse: $P_{\mathfrak{B}\leftarrow S} =  P^{-1}_{S\leftarrow \mathfrak{B}}$.

Whichever way you decide to do this make sure you save your result as `PBS`.

<!-- grade_id: cell-ba87a12dee214f26, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-7e03f3eac2c68ea7, points: 1, grade: true, solution: false -->
```python
"Verify PBS is a NumPy array of size (4,4). (1 mark)"
assert isinstance(PBS,np.ndarray) and PBS.shape == (4,4)
print("Test 1a.1: Success!")
```

<!-- grade_id: cell-437b9e0538462163, points: 1, grade: true, solution: false -->
```python
"Verify the entries in the first column of PBS are correct to one decimal place. (1 mark)"
assert np.allclose(np.round(PBS[:,0],1),[-1.1, -0.3,  0.4,  0.4])
print("Test 1a.2: Success!")
```

<!-- grade_id: cell-b1bbf6588929a510, points: 1, grade: true, solution: false -->
```python
"Verify matrix PBS has correct entries. This cell contains hidden tests. (1 mark)"
```

## <span style="color:#87CEEB">Problem 1b:</span>
For $\vec{x}=\begin{bmatrix}1\\2\\3\\4\end{bmatrix}$, use the matrix $P_{\mathfrak{B}\leftarrow S}$ from part (a) to find $[\vec{x}]_\mathfrak{B}$. Save your answer as `xB`.

<!-- grade_id: cell-9851dcd97cd7a505, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-62ab46c2dddb3d2c, points: 1, grade: true, solution: false -->
```python
"Verify xB is a NumPy array of size (4,). (1 mark)"
assert isinstance(xB,np.ndarray) and xB.shape == (4,)
print("Test 1b.1: Success!")
```

<!-- grade_id: cell-5ba144bc35b1da55, points: 1, grade: true, solution: false -->
```python
"Verify the first entry in xB is correct to one decimal place. (1 mark)"
assert np.allclose(np.round(xB[0],1),-1.2)
print("Test 1b.2: Success!")
```

<!-- grade_id: cell-2d06a4859744523f, points: 1, grade: true, solution: false -->
```python
"Verify vector xB has correct entries. This cell contains hidden tests. (1 mark)"
```

## <span style="color:#87CEEB">Problem 1c:</span>
Let $L: \mathbb{R}^4\rightarrow \mathbb{R}^4$ be the linear mapping:

$L(x_1, x_2, x_3, x_4) = (x_2+5x_4,\quad x_1-x_3,\quad x_1+x_2+x_3+x_4,\quad 0)$

Calculate the matrix $[L]_\mathfrak{B}$. Save the matrix with the name `LB`.

**Hint:** Use the fact that $[L]_\mathfrak{B} = P_{\mathfrak{B}\leftarrow S} [L]_S P_{S\leftarrow \mathfrak{B}} = P_{\mathfrak{B}\leftarrow S} [L]_S P^{-1}_{\mathfrak{B}\leftarrow S}$.

<!-- grade_id: cell-9f835ce501dde37e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-b673b3dd31fe038b, points: 1, grade: true, solution: false -->
```python
"Verify LB is a NumPy array of size (4,4). (1 mark)"
assert isinstance(LB,np.ndarray) and LB.shape == (4,4)
print("Test 1c.1: Success!")
```

<!-- grade_id: cell-07c45dfd4cb78259, points: 1, grade: true, solution: false -->
```python
"Verify the entries in the first column of LB are correct to one decimal place. (1 mark)"
assert np.allclose(np.round(LB[:,0],1),np.array([-5. , -3.2,  2.9,  3.4]))
print("Test 1c.2: Success!")
```

<!-- grade_id: cell-3664f6351cbca50b, points: 1, grade: true, solution: false -->
```python
"Verify matrix LB has correct entries. This cell contains hidden tests. (1 mark)"
```

## <span style="color:#87CEEB">Problem 1d:</span>
For the vector $[\vec{x}]_\mathfrak{B}$ that you found in part (a), use the matrix in (c) to calculate the vector $[L(\vec{x})]_\mathfrak{B}$. Save the vector as `LxB`.

<!-- grade_id: cell-9ea158daf534243b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-4e80ccd31f2c98af, points: 1, grade: true, solution: false -->
```python
"Verify LxB is a NumPy array of size (4,). (1 mark)"
assert isinstance(LxB,np.ndarray) and xB.shape == (4,)
print("Test 1d.1: Success!")
```

<!-- grade_id: cell-15fe106779baa7c4, points: 1, grade: true, solution: false -->
```python
"Verify the first entry in LxB is correct to one decimal place. (1 mark)"
assert np.allclose(np.round(LxB[0],1),-19.3)
print("Test 1d.2: Success!")
```

<!-- grade_id: cell-0dfb94eb6e3c1575, points: 1, grade: true, solution: false -->
```python
"Verify vector LxB has correct entries. This cell contains hidden tests. (1 mark)"
```

## <span style="color:#87CEEB">Problem 2: Power Method</span>

Let $A=\left[\begin{array}{cccc} 1&2&0&0 \\ 1&1&1&1\\0&1&1&5\\0&0&1&1\end{array}\right]$. We will use the Power Method to approximate the dominant eigenvalue of $A$, and a dominant eigenvector.

## <span style="color:#87CEEB">Problem 2a:</span>

First choose a nonzero initial vector $\vec{x}_0$ that has a length of $1$ - we'll choose $\vec{x}_0 = \begin{bmatrix}1\\0\\0\\0\end{bmatrix}$. 

Next, compute $\vec{x}_1, \vec{x}_2, \vec{x}_3$ based on the Power Method (that is, $\vec{x}_1=A\vec{x}_0/\|A\vec{x}_0\|$, etc., so the vectors are normalized to length $1$). The first two are done for you below, add another couple lines of code to compute $\vec{x}_3$. Save the vector with the name `x3`.

<!-- grade_id: cell-c05d5733be2b52c3, points: 0, grade: false, solution: true -->
```python
A = np.array([[1,2,0,0],
              [1,1,1,1],
              [0,1,1,5],
              [0,0,1,1]])

# our initial unit vector
x0 = np.array([1,0,0,0]) 

y1 = A@x0                         # compute next vector y1
x1 = y1/(np.sqrt(np.dot(y1,y1)))  # normalize the vector y1 to form x1
print(f"x1 = {x1}")

y2 = A@x1                         # compute next vector y2
x2 = y2/(np.sqrt(np.dot(y2,y2)))  # normalize the vector y2 to form x2
print(f"x2 = {x2}")

# YOUR CODE HERE


```

<!-- grade_id: cell-07042cb51a276c18, points: 1, grade: true, solution: false -->
```python
"Verify x3 is a NumPy array of size (4,). (1 mark)"
assert isinstance(x3,np.ndarray) and x3.shape == (4,)
print("Test 2a.1: Success!")
```

<!-- grade_id: cell-cd4e4c783aa1e1a8, points: 1, grade: true, solution: false -->
```python
"Verify the first entry in x3 is correct to 8 decimal places. (1 mark)"
assert np.allclose(np.round(x3[0],8),0.71818485)
print("Test 2a.2: Success!")
```

<!-- grade_id: cell-15325865b4f9255a, points: 1, grade: true, solution: false -->
```python
"Verify vector x3 has correct entries to 8 decimal places. This cell contains hidden tests. (1 mark)"
```

## <span style="color:#87CEEB">Problem 2b:</span>

Now we'll compute $\vec{x}_4, \vec{x}_5, \ldots$ (continuing past $\vec{x}_3$), until the results **converges** to a dominant eigenvector of $A$. <br />
**Note**: We'll consider convergence to mean that the *last two* computed vectors agree to at least 4 decimal places in each component. 

Most of the code is already implemented below, **just change the upper limit in the `range` command** so that the loop runs longer. Currently the upper bound in the loop is set to **3**, change this to a larger number and inspect the output until you see each of the four components stabilize at 4 decimal places.

**Save the dominant eigenvector** with the name `x_d` (don't round, just truncate to 4 decimal places). You can just create an array named `x_d` and type in the entries stopping after the 4th decimal digit for each component.

**Save the corresponding dominant eigenvalue** as `lambda_d` and truncate at 2 decimal places. You can `lambda_d = ` and give the value by stopping after the 2nd decimal digit.

<!-- grade_id: cell-15d94f0b6796bb02, points: 0, grade: false, solution: true -->
```python
x0 = np.array([1,0,0,0])
evec_approx_list = [x0]  # this will be the list of approximations to the dominant eigenvector
eval_approx_list = [np.dot(A@x0,x0)]  # this will be the list of approximations to the dominant eigenvalue

# ONLY CHANGE THE UPPER BOUND IN `range`. MAKE IT BIG ENOUGH SO TWO CONSECUTIVE VECTORS AGREE TO 4 DECIMAL PLACES.
for i in range(1,3):  
    new_vec = A@evec_approx_list[-1]                                 # pick last vector in evec_approx_list, and apply A to it
    evec_approx_list.append(new_vec/np.sqrt(np.dot(new_vec,new_vec)))            # add the new normalized vector to evec_approx_list
    eval_approx_list.append(np.dot(A@evec_approx_list[-1],evec_approx_list[-1])) # add updated approx. for lambda to eval_approx_list  
    print(f"x_{i:<2} = {evec_approx_list[i]}, and λ^({i:<2}) = {eval_approx_list[i]}")  # print the new approximations to the screen

# Now enter your answers for x_d and lambda_d below after you inspected the output.
# YOUR CODE HERE

```

<!-- grade_id: cell-3bbecd728ade51f6, points: 1, grade: true, solution: false -->
```python
"Verify x_d is a NumPy array of size (4,). (1 mark)"
assert isinstance(x_d,np.ndarray) and x_d.shape == (4,)
print("Test 2b.1: Success!")
```

<!-- grade_id: cell-f0be6061777770d9, points: 1, grade: true, solution: false -->
```python
"Verify the first entry in x_d is correct. (1 mark)"
assert np.allclose(x_d[0],0.3953)
print("Test 2b.2: Success!")
```

<!-- grade_id: cell-fc024e3134a5ab68, points: 1, grade: true, solution: false -->
```python
"Verify lambda_d is a float. (1 mark)"
assert isinstance(lambda_d,float)
print("Test 2b.3: Success!")
```

<!-- grade_id: cell-78d9b9c3e976b65d, points: 1, grade: true, solution: false -->
```python
"Verify vector x_d has correct entries to 4 decimal places. This cell contains hidden tests. (1 mark)"
```

<!-- grade_id: cell-82c7f4fbd1305944, points: 1, grade: true, solution: false -->
```python
"Verify lambda_d is correct to 2 decimal places. (1 mark)"
```

```python

```