# MATH 223 Computing Assignment 2

Load the necessary packages:

```python
import numpy as np
import matplotlib.pyplot as plt  # used for plotting points in a graph
```

## Linear Transformations in 2D

### Import Dataset

Import the dataset `points.csv`:

```python
points = np.loadtxt('points.csv',delimiter=',')
```

The dataset `points.csv` is a 2 by 1000 matrix:

```python
points.shape
```

Each column in `points` represents a point in $\mathbb{R}^2$. Create a scatter plot to visualize the points. In the cell below, the parameter `s = 2` specifies the size of the points in the plot and `c = [[0.75,0,0,1]]` specifies the color as an RGBA vector `[red,green,blue,opacity]`.

```python
plt.scatter(points[0,:],points[1,:],s=2,c=[[0.75,0,0,1]])
plt.axis('equal'), plt.grid(True)
plt.show()
```

Now let's apply some geometrical transformations to the points in the data set, and plot the result data.

### Rotation Matrix

A rotation matrix is given by

$$
R_{\theta} = \begin{bmatrix} \cos(\theta) & -\sin(\theta) \\ \sin(\theta) & \cos(\theta) \end{bmatrix}
$$

The linear transformation corresponding to $R_{\theta}$ is rotation counterclockwise by $\theta$ radians.

```python
# construct the rotation matrix:
theta = np.pi/6
R = np.array([[np.cos(theta),-np.sin(theta)],
              [np.sin(theta),np.cos(theta)]])

# apply the matrix to the points with matrix multiplication:
points1 = R@points

# plot the collection of points (you don't need to worry about the syntax below,
# this will already be filled out in the susequent questions. You'll just focus
# on constructing the correct transformation matrix.
plt.scatter(points[0,:],points[1,:],s=2,c=[[0.75,0,0,0.1]]) # original points, dimmed
plt.scatter(points1[0,:],points1[1,:],s=2,c=[[0.75,0,0,1]]) # transformed points
plt.axis('equal'), plt.grid(True)
plt.show()
```

### Scaling Matrix

A scaling matrix is given by

$$
C_{a,b} = \begin{bmatrix} a & 0 \\ 0 & b \end{bmatrix}
$$

The linear transformation corresponding to $C_{a,b}$ is scaling by $a$ is the $x$-direction and scaling by $b$ in the $y$-direction.

```python
# construct the scaling matrix:
a = 3
b = 1/2
C = np.diag([a,b]) # this creates a 2x2 diagonal matrix with entries a,b along the diagonal

# apply the matrix to the points with matrix multiplication:
points2 = C@points

# plot the collection of points
plt.scatter(points[0,:],points[1,:],s=2,c=[[0.75,0,0,0.05]]) # original points, dimmed
plt.scatter(points2[0,:],points2[1,:],s=2,c=[[0.75,0,0,1]]) # transformed points
plt.axis('equal'), plt.grid(True)
plt.show()
```

### Shearing Matrix

A shear matrix is given by

$$
S_u = \begin{bmatrix} 1 & u \\ 0 & 1 \end{bmatrix}
$$

The linear transformation corresponding to $S_u$ is shearing by factor $u$ in the $x$-direction.

```python
# construct the scaling matrix:
u = 1
S = np.array([[1,u],
              [0,1]])

# apply the matrix to the points with matrix multiplication:
points3 = S@points

# plot the collection of points
plt.scatter(points[0,:],points[1,:],s=2,c=[[0.75,0,0,0.1]]) # original points, dimmed
plt.scatter(points3[0,:],points3[1,:],s=2,c=[[0.75,0,0,1]]) # transformed points
plt.axis('equal'), plt.grid(True)
plt.show()
```

### Reflection Matrix

A reflection matrix is given by

$$
F_x = \begin{bmatrix} 1 & 0 \\ 0 & -1 \end{bmatrix}
$$

The linear transformation corresponding to $F_x$ is a reflection across the $y$-axis.

Similarly, 

$$
F_y = \begin{bmatrix} -1 & 0 \\ 0 & 1 \end{bmatrix}
$$

is a reflection across the $y$-axis.

```python
# construct the reflection matrix:
F = np.array([[1,0],
              [0,-1]])

# apply the matrix to the points with matrix multiplication:
points4 = F@points

# plot the collection of points
plt.scatter(points[0,:],points[1,:],s=2,c=[[0.75,0,0,0.1]]) # original points, dimmed
plt.scatter(points4[0,:],points4[1,:],s=2,c=[[0.75,0,0,1]]) # transformed points
plt.axis('equal'), plt.grid(True)
plt.show()
```

### Combining Transformations with Matrix Multiplication

We can combine linear transformations using matrix multiplication. Recall the transformations in the product are applied from right-to-left, meaning that in the product below $R$ is applied first, then $C$, and finally $S$.

```python
# construct the scaling matrix:
points5 = R @ S @ C @ points

# plot the collection of points
plt.scatter(points[0,:],points[1,:],s=2,c=[[0.75,0,0,0.05]])
plt.scatter(points5[0,:],points5[1,:],s=2,c=[[0.75,0,0,1]])
plt.axis('equal'), plt.grid(True)
plt.show()
```

## <span style="color:#fdb73e">Practice Problem</span>

A linear transformation was applied to the point to produce the following image, find the matrix for this linear transformation.

![prob1a.png](prob_practice.png)

<span style="color:#fdb73e">Solution:</span>: Since the little circle is now on the opposite side of the heart this means the image has been reflected. Let's start by refecting across the y-axis.

Try this for yourself in the code box below!

This produces the image:

![prob_practice1.png](prob_practice1.png)

Now it looks like the final image is a vertical shear. By how much? Try experimenting with some values. Looks like 1.5 does the trick.


<!-- grade_id: cell-72d7d316da19f493, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE


# construct some matrices that might contribute to the final image
F = np.array([[-1,0],
              [0,1]])  # reflect across y-axis
S = np.array([[1,0],
             [1.5,1]]) # vertical shear by a factor of 1.5
T = S@F

# apply the matrix product to the collection of points
points6 = T@points

# plot the collection of points
plt.scatter(points[0,:],points[1,:],s=2,c=[[0.75,0,0,0.05]])
plt.scatter(points6[0,:],points6[1,:],s=2,c=[[0.75,0,0,1]])
plt.axis('equal'), plt.grid(True)
plt.show()
```

With the practice problem done, try this out for yourself in the following two questions.

## <span style="color:#87CEEB">Problem 1a:</span>

Find the matrix corresponding to the linear transformation in the image below. The matrix is a combinations of the basic geometrical transformations. Save the matrix as `T1a`.

![prob1a.png](prob1a.png)

Begin your solution by building the matrix `T1a` as a product of some of the basic transformation matrices:

<!-- grade_id: cell-a2000369afa93781, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE


# apply the matrix to the points with matrix multiplication:
points7 = T1a @ points

# plot the collection of points
plt.scatter(points[0,:],points[1,:],s=2,c=[[0.75,0,0,0.1]])
plt.scatter(points7[0,:],points7[1,:],s=2,c=[[0.75,0,0,1]])
plt.axis('equal'), plt.grid(True)
plt.show()
```

<!-- grade_id: cell-32079494eae50864, points: 1, grade: true, solution: false -->
```python
"Verify T1a is a NumPy array of size (2,2). (1 mark)"
assert isinstance(T1a,np.ndarray)
assert T1a.shape == (2,2)
print("Problem 1a Test 1: Success!")
```

<!-- grade_id: cell-b4f543c302fee3e7, points: 1, grade: true, solution: false -->
```python
"Verify T1a has correct entry in upper left corner to 1 decimal places. (1 marks)"
assert np.round(T1a[0,0],1) == 0.7 ,  "Check entries to one decimal place."
print("Problem 1a Test 2: Success!")
```

<!-- grade_id: cell-fa602033be0ef9b4, points: 2, grade: true, solution: false -->
```python
"Verify T1a has correct entries. This cell contains hidden tests. (2 marks)"
```

## <span style="color:#87CEEB">Problem 1b:</span>

Find the matrix corresponding to the linear transformation in the image below. The matrix is a combinations of the basic geometrical transformations. Save the matrix as `T1b`.

![prob1b.png](prob1b.png)

Begin your solution by building the matrix `T1b` as a product of some of the basic transformation matrices:

<!-- grade_id: cell-dc9b084d2e268a0c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE


# apply the matrix to the points with matrix multiplication:
points8 = T1b @ points

# plot the collection of points
plt.scatter(points[0,:],points[1,:],s=2,c=[[0.75,0,0,0.1]])
plt.scatter(points8[0,:],points8[1,:],s=2,c=[[0.75,0,0,1]])
plt.axis('equal'), plt.grid(True)
plt.show()

```

<!-- grade_id: cell-c4609fdb79b61241, points: 1, grade: true, solution: false -->
```python
"Verify T1b is a NumPy array of size (2,2). (1 mark)"
assert isinstance(T1b,np.ndarray)
assert T1a.shape == (2,2)
print("Problem 1b Test 1: Success!")
```

<!-- grade_id: cell-2f36dbdcc2f59a60, points: 1, grade: true, solution: false -->
```python
"Verify T1b has entries that sum up to 1.5. (1 marks)"
assert np.allclose(T1b.sum(),1.5) , "Entries should add up to 1.5."
print("Problem 1b Test 2: Success!")
```

<!-- grade_id: cell-bda410b2ff74e94f, points: 2, grade: true, solution: false -->
```python
"Verify T1b has correct entries. This cell contains hidden tests. (2 marks)"
```

```python

```