# Math 232 Computing Assignment 1

## Basic Python Commands for Linear Algebra:

Let's learn some commands to get up and running with Python. 

First we start by importing the numerical python (numpy) package that contains some basic tools for doing mathematics in Python. We import it to the abreviated name `np`.

```python
import numpy as np
```

### Matrices and Vectors:

To enter a matrix (which includes column and row vectors) we use np.array and pass a list of rows:

```python
A = np.array([[1,2,3],[4,5,6],[7,8,9]])  # this is a matrix with three rows: [1,2,3], [4,5,6], [7,8,9]
print(A)  # we can print the matrix to the screen for viewing
```

```python
v = np.array([1,-1,1]) # this is a vector 
print(v)
```

To form an augmented matrix $[A|v]$ you could either enter it as a whole matrix from scratch, or if $A$ and $v$ are already created then you can combine them into one matrix as follows:

```python
aug_mat = np.column_stack([A,v])
print(aug_mat)
```

### RREF:

To apply Gaussian-Elimination to a matrix you'll need the `rref` package that we wrote for the course. This is because numpy doesn't have a built in `rref` command.
To load the package make sure you have the `rref.py` file downloaded from Canvas and then uploaded into syzygy next to your assignment file. You should see it in the file browser window on the left. We load it with the following command:

```python
from rref import *
```

Now we can use it to do Gaussian elimination. We'll have python reduced the `aug_mat` (which we defined above) to reduced row echelon form.

```python
rref(aug_mat)
```

The `rref` command returns the reduced row-echelon form of the matrix, along with a list of pivot columns (numbering starts at 0, which means the first column is called 'column 0', the second is 'column 1', etc.).

### Rank:

To compute the rank of a matrix we can use the command `np.linalg.matrix_rank`. Next, we compute the rank of matrix `A`.

```python
np.linalg.matrix_rank(A)
```

This returned the integer $2$ (but it is also telling us that it was stored as a 64-bit float... 'how' an integer is stored in computer memory is something we won't be concerned with).

### Matrix Multiplication:

To multiply matrices use the `@` operator. For example, to compute $A\vec{v}$ we do:

```python
A@v  # matrix-vector product
```

```python
A@A  # matrix-matrix product
```

## A Worked Example:

Now we are ready to solve a system of linear equations. Consider the system <br />

$x_1 + 2x_2 + 3x_3 = 4$ <br />
$4x_1 + 5x_2 + 6x_3 = -2$

First we form the augmented matrix. To do this we will build the coefficient matrix, and the vector for the constants on the right side of the equations, then use `column_stack` to augment them. Finally we call `rref`:

```python
A = np.array([[1,2,3],
              [4,5,6]])
b = np.array([4,-2])
aug_mat = np.column_stack([A,b])
print(aug_mat)
rref(aug_mat)
```

The result of this calcluation tells us there is one free variable, and the general solution is:

$ \vec{x} = \begin{bmatrix} -8 \\ 6 \\ 0\end{bmatrix} + t\begin{bmatrix} 1 \\ -2\\ 1\end{bmatrix}, \qquad t\in \mathbb{R}$

Notice that $\begin{bmatrix} -8 \\ 6 \\ 0\end{bmatrix}$ is one particular solution to the linear system, whereas $\begin{bmatrix} 1 \\ -2\\ 1\end{bmatrix}$ is a solution to the homogeneous system. We can check this with multiplication:

```python
A@np.array([-8,6,0])  # this should return [4,-2] since that is the right-hand side of the linear system.
```

```python
A@np.array([1,-2,1]) # this should return [0,0]
```

Finally, if we were just interested in know whether there was a solution, and if so, how many free variables, then we could have computed the rank:

```python
print(np.linalg.matrix_rank(A))  # rank of coefficient matrix
print(np.linalg.matrix_rank(aug_mat))  # rank of augmented matrix
```

The ranks are equal, so the rightmost column is not a pivot column which means the system is consistent. 
With 3 variables, and a rank of 2, there must be 1 free variable. 

## Assignment Exercises:

```python
# import the necessary packages (not really necessary as we already did it above - but placed here for convenience)
import numpy as np
from rref import *
```

Output:
```
rref package loaded successfully
```

**Question 1:** Consider the following vectors in $\mathbb{R}^5$:

   $$\vec{v}_0 = \begin{bmatrix}1\\1\\2\\3\\-1\end{bmatrix}, 
   \vec{v}_1 = \begin{bmatrix}1\\2\\2\\1\\0\end{bmatrix}, 
   \vec{v}_2 = \begin{bmatrix}1\\4\\2\\-3\\2\end{bmatrix}, 
   \vec{v}_3 = \begin{bmatrix}1\\7\\2\\-9\\5\end{bmatrix},
   \vec{v}_4 = \begin{bmatrix}2\\1\\2\\1\\3\end{bmatrix},
   \vec{v}_5 = \begin{bmatrix}5\\0\\4\\4\\8\end{bmatrix}$$

**(a)** Create each of the 6 vectors in python as an `np.array`, and save them with the names `v0` through `v5`.<br />
Then, use the command `np.column_stack` to create the matrix with column vectors `v0` through `v5`. Save it with the name `A`.

<!-- grade_id: cell-58dc62a1aeb504c5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-b77514798d80ab40, points: 1, grade: true, solution: false -->
```python
"Verify v0 through v5 were created as arrays and have the correct size. (1 mark)"
assert isinstance(v0,np.ndarray) and v0.shape == (5,)
assert isinstance(v1,np.ndarray) and v1.shape == (5,)
assert isinstance(v2,np.ndarray) and v2.shape == (5,)
assert isinstance(v3,np.ndarray) and v3.shape == (5,)
assert isinstance(v4,np.ndarray) and v4.shape == (5,)
assert isinstance(v5,np.ndarray) and v5.shape == (5,)
print("Test 1: Success!")
```

<!-- grade_id: cell-3e5d8306d983f75c, points: 1, grade: true, solution: false -->
```python
"Verify matrix A is an array of the correct size. (1 mark)"
assert isinstance(A,np.ndarray)
assert A.shape == (5,6)
print("Test 2: Success!")
```

<!-- grade_id: cell-27aa1b967690a4f1, points: 1, grade: true, solution: false -->
```python
"Verify vector v0 has the correct entries. (1 mark)"
assert np.array_equal(v0, np.array([1,1,2,3,-1]))
print("Test 3: Success!")
```

<!-- grade_id: cell-6000bf6fd9b445e9, points: 0, grade: true, solution: false -->
```python
"Verify the first and second column of A are correct. (1 mark)"
assert np.array_equal(A[:,0], np.array([1,1,2,3,-1]))
assert np.array_equal(A[:,1], np.array([1,2,2,1,0]))
print("Test 4: Success!")
```

<!-- grade_id: cell-59019cf5ae6678ec, points: 1, grade: true, solution: false -->
```python
"Verify vectors v1 through v5 have correct entries. This cell contains hidden tests. (1 mark)"
```

<!-- grade_id: cell-20bf8bd01ba72f83, points: 1, grade: true, solution: false -->
```python
"Verify matrix A has correct entries. This cell contains hidden tests. (1 mark)"
```

**(b)** Using the `rref` command, find the reduced row echelon form of `A`. Save the output of the `rref` command as `RA`.

<!-- grade_id: cell-d58e8ed47f442fb7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-55f740b0bba7d4ea, points: 1, grade: true, solution: false -->
```python
"Verify RA was computed from the rref command. (1 mark)"
assert isinstance(RA[0],np.ndarray)
assert isinstance(RA[1],tuple)
assert RA[0].shape == (5,6)
print("Test 5: Success!")
```

**(c)** Which columns are the pivot columns? Give your answer in the form of a tuple of column indices, the same form that was output from the `rref` command: `( , , )`. Save your answer as `pivot_indices`. <br />
Hint: you should type in the cell below `pivot_indices = ( , , )`, with the three correct index numbers separated by the commas.

<!-- grade_id: cell-41a0433764cf8794, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-9c2195c1eea6495f, points: 1, grade: true, solution: false -->
```python
"Verify pivot_indices is a tuple of the correct size. (1 mark)"
assert isinstance(pivot_indices,tuple)
assert len(pivot_indices) == 3
print("Test 6: Success!")
```

<!-- grade_id: cell-004c749613ad8b05, points: 1, grade: true, solution: false -->
```python
"Verify all entries of pivot_indices are correct. This cell contains hidden tests. (1 mark)"
```

**(d)** What is the rank of matrix `A`? Save it with the name `rankA`.

<!-- grade_id: cell-71f0735fc50a2f77, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-22389fd4a38df1c7, points: 1, grade: true, solution: false -->
```python
"Verify rank A is correct. This cell contains hidden tests. (1 mark)"
```

**(e)** How many free variables are there in the system $A\vec{x} = \vec{0}$? Save your answer as `nullityA`.

<!-- grade_id: cell-1405abe074d45e37, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-f14b63044200fc63, points: 1, grade: true, solution: false -->
```python
"Verify number of free variables, nullityA, is correct. This cell contains hidden tests. (1 mark)"
```

**(f)** Create a new matrix containing only the pivot columns of `A`. Save this matrix with the name `P`.

<!-- grade_id: cell-d8a4491cdb7d715e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

```python
"Verify pivot columns matrix P is np.array and has correct correct size. (1 mark)"
assert isinstance(P,np.ndarray)
assert P.shape == (5,3)
print("Test 7: Success!")
```

<!-- grade_id: cell-57eb2cf0b32e6fb8, points: 1, grade: true, solution: false -->
```python
"Verify pivot columns matrix P has correct entries. This cell contains hidden tests. (1 mark)"
```

**(g)**  Find a vector $\vec{w}$ so that the product $P \vec{w} = \vec{v}_2$. Use an `np.array` to create `w` and save it with the name `w`.  

<!-- grade_id: cell-6c1f12ec782eb0b1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-f77b688c282fc917, points: 1, grade: true, solution: false -->
```python
"Verify w is a vector of the correct size. (1 mark)"
assert isinstance(w,np.ndarray)
assert w.shape == (3,)
print("Test 8: Success!")
```

<!-- grade_id: cell-95b36754da59d694, points: 1, grade: true, solution: false -->
```python
"Verify w has the correct entries. This cell contains hidden tests. (1 mark)"
```

**(h)** Find a matrix $W$ so that $PW = [\vec{v}_2 | \vec{v}_3 | \vec{v}_5 ]$. Use an `np.array` to create $W$ and save it with the name `W`.

<!-- grade_id: cell-44fd1385c7cbc20a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-f0375d82e625506b, points: 1, grade: true, solution: false -->
```python
"Verify W is a matrix of the correct size. (1 mark)"
assert isinstance(W,np.ndarray)  # checking if W was constructed as an array
assert W.dtype == int   # checking if the entries in W are integers (no decimal points at all)
assert W.shape == (3,3) # checking if W is a 3x3 matrix
print("Test 8: Success!")
```

<!-- grade_id: cell-102416f0c9aea6b3, points: 1, grade: true, solution: false -->
```python
"Verify W has the correct entries. This cell contains hidden tests. (1 mark)"
```

```python

```